<img src='https://hammondm.github.io/hltlogo1.png' style="float:right">

Linguistics 578<br>
Fall 2023<br>
Hammond

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Please do not add, remove, or copy autograded cells.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

***

***my name***: [put your name here]

***people I talked to about the assignment***: [put your answer here]

***

## Homework #6

Here are the imports. Please do not import anything else.

```python
import os,re
import numpy as np
from scipy.io import wavfile
import matplotlib.pyplot as plt
from IPython.display import Audio
```

1. We need to use the `kal` database, so here's a question that just lets you specify where it is and then checks that. (Please leave the first two comment lines in place and just add your own below.)

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
#kal = '/mhdata/cmu_us_kal_diphone/'
#kal = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
assert len(os.listdir(kal)) == 20
```

We need these from this week's notebook. I've converted the code that creates the diphone database into a function. Make sure the run the cell below before going on.

```python
def getbit(start,end,filename,dirname=kal):
    sr,w = wavfile.read(dirname+'wav/'+filename)
    startidx = int(sr * start)
    endidx = int(sr*end)
    return w[startidx:endidx]

def makeDiphones(where):
    diphones = {}
    for i in range(1,1350):
        filename = str(i)
        while len(filename) < 4: filename = '0' + filename
        filename = 'kal_' + filename + '.lab'
        f = open(kal+'lab/'+filename,'r')
        t = f.read()
        f.close()
        t = re.sub(' 26 ','',t)
        t = t.split('\n')
        t = t[3:-1]
        fset = []
        start = 0
        for line in t:
            bits = line.split('\t')
            bits = bits[1:]
            fset.append((float(start),float(bits[0]),bits[1]))
            start = bits[0]
        i = 1
        while i < len(fset):
            diphone = fset[i-1][2] + '-' + fset[i][2]
            start = (fset[i-1][1] - fset[i-1][0])/2 + fset[i-1][0]
            end = (fset[i][1] - fset[i][0])/2 + fset[i][0]
            filename = re.sub('\.lab$','.wav',filename)
            triple = (filename,start,end)
            if diphone in diphones:
                diphones[diphone].append(triple)
            else:
                diphones[diphone] = [triple]
            i += 1
    return diphones

diphones = makeDiphones(kal)
```

2. For this question, you need to download and unzip the CMU pronouncing dictionary from [here](https://github.com/Alexir/CMUdict). (Again, do not remove the first two lines below; just add your own below.)

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
#cmufile = '/mhdata/CMUdict-master/cmudict-0.7b'
#cmufile = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
f = open(cmufile,'rb')
t = f.read()
f.close()
cmudict = t.decode('ISO-8859-1')
assert np.isclose(len(cmudict),3865710,atol=2)
```

3. Your task is to parse the `cmudict` variable into a dictionary from words to letter sequences. Be mindful of the following:
  1. The file format is funky.
  2. There may be errors and irregularities to recover from.
  3. You have to remove the header lines and final empty line.
  4. You have to convert to lower case.
  5. You have to remove stress information.
  6. You have to accommodate the fact that some words have multiple pronunciations.

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def makeDictionary(datafile):
    '''map the cmudict text file to a dictionary
    args:
        datafile: the cmudict07b file as a list of lines
    returns:
        a dictionary from spellings to a list of lists of
        characters.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
d = makeDictionary(cmudict)
assert len(d) == 125770
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
assert d['tomato'] == [
    ['t', 'ah', 'm', 'ey', 't', 'ow'],
    ['t', 'ah', 'm', 'aa', 't', 'ow']
]
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
assert d["didn't"] == [
    ['d', 'ih', 'd', 'ah', 'n', 't'],
    ['d', 'ih', 'd', 'n', 't'],
    ['d', 'ih', 'd', 'ah', 'n'],
    ['d', 'ih', 'n', 't']
]
```

4. Now write a function that takes a sentence, breaks it into words, looks up those words, converts the transcription to diphones, and returns a list of diphones. In this exercise, we'll keep word transitions simple and separate all words with `pau`. (If a word is not in the dictionary, the function should return *uh* in its stead.)

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
def sentence2diphones(s,c):
    '''converts a sentence into a sequence of diphones
    args:
        s: a sentence as a string
        c: cmu dictionary as a dictionary
    returns:
        a list of diphones with 'pau' at the beginning
        and end, and between words
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
words = "I'm feeling pretty optimistic right now."
res = sentence2diphones(words,d)
assert len(res) == 33
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
assert res == [
    'pau-ay','ay-m','m-pau','pau-f','f-iy',
    'iy-l','l-ih','ih-ng','ng-pau','pau-p',
    'p-r','r-ih','ih-t','t-iy','iy-pau',
    'pau-aa','aa-p','p-t','t-ih','ih-m',
    'm-ih','ih-s','s-t','t-ih','ih-k',
    'k-pau','pau-r','r-ay','ay-t','t-pau',
    'pau-n','n-aw','aw-pau'
]
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
words = "The bznork sings!"
res = sentence2diphones(words,d)
assert res == [
    'pau-dh','dh-ah','ah-pau','pau-ah','ah-pau',
    'pau-s','s-ih','ih-ng','ng-z','z-pau'
]
```

5. Finally, let's write a function that takes a list of diphones and creates a wave file.

<!-- grade_id: q5a, points: 0, grade: false, solution: true -->
```python
def makeWave(s,dps):
    '''converts a sequence of diphones into a wave
    args:
        s: a sequence of diphones
        dps: diphone database
    returns:
        a wave as a numpy array
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q5t1, points: 1, grade: true, solution: false -->
```python
s = "I like pancakes."
res = sentence2diphones(s,d)
w = makeWave(res,diphones)
assert type(w) == np.ndarray
```

<!-- grade_id: q5t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(len(w),58119,atol=50)
```

And now, just for fun, you can look at and listen to what you've made:

```python
plt.plot(w)
plt.show()
```

```python
Audio(w,rate=16000)
```